Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
StudentID = ""
```

---

# COS760 2025 - HW1-1 (20 points)

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)


Author(s): Prof Vukosi Marivate, Fiskani Banda

**Note:** *This notebooks is only for students taking COS 760. It should not be shared outside the class without the permission of the lecturer(s).*



**NLP Fundamentals and Text Preprocessing – Homework Assignment**  

Welcome to your first **NLP homework**! This assignment will test your understanding of **Natural Language Processing (NLP)** concepts and challenge you with practical text preprocessing tasks.  

- Apply text preprocessing techniques using **NLTK and regular expressions** to clean and analyze text. You'll convert text to lowercase, remove punctuation, and filter out stopwords—key steps in real-world NLP tasks.  
- 🔍 **Bonus Challenge:** Compare NLTK’s approach with sklearn’s preprocessing—spot the differences and refine your understanding of text cleaning!  

By completing this assignment, you’ll **build essential NLP skills** that will help you tackle real-world language challenges. Ready to dive in? 🚀

#### Resources

You might want to get acquianted with the online free edition of the great book "Natural Language Processing with Python" [NLTK Book](https://www.nltk.org/book/)

---
## Get Comfortable with Google Colab

Execute a simple Python cell in Colab.

**📌 Instruction:** Execute a simple Python cell in Colab.

```python
# Run this cell
print("Hello, COS760 NLP!")
```

**📌 Instruction:** Run shell commands in Colab.

```python
!ls  # List files
!pip list  # Show installed Python packages
!cat /proc/cpuinfo | grep "model name"  # Check Colab CPU
```

---
## Main Assignment

Using Python, split the following text into words and sentences manually and using nltk.tokenize. Compare the results.

```python
import nltk
from sklearn.feature_extraction.text import CountVectorizer
nltk.download('punkt')

text = "Prof. Marivate and Dr. Modupe teach NLP at UP. The University of Pretoria (UP) is in South Africa. NLP is challenging! However, many researchers and companies are working to solve its problems."

```

1. Tokenize the given text using the function**[nltk.word_tokenize](https://www.nltk.org/api/nltk.tokenize.html)** and save the tokenized text in variable`nltk_tokens`

<!-- grade_id: cell-f16aa844a5d4b143, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print("NLTK Tokens:", nltk_tokens)
```

<!-- grade_id: cell-b57a35a80a11a4ed, points: 1, grade: true, solution: false -->
```python
assert isinstance(nltk_tokens, list), "NLTK tokenization failed." # 1 point
```

<!-- grade_id: cell-cb078cc15b92cda7, points: 2, grade: true, solution: false -->
```python
assert len(nltk_tokens) == 39, "NLTK token list is empty." # 2 point
```

2. Tokenize the given text using **[CountVectorizer](https://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.CountVectorizer.html)** and **[CountVectorizer.get_feature_names_out()](https://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.CountVectorizer.html#sklearn.feature_extraction.text.CountVectorizer.get_feature_names_out)** in variable`sklearn_tokens`

<!-- grade_id: cell-114b1969a640cd58, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print("Scikit-learn Tokens:", sklearn_tokens)
```

<!-- grade_id: cell-91ef45333f6f10b7, points: 1, grade: true, solution: false -->
```python
assert isinstance(sklearn_tokens, list), "Scikit-learn tokenization failed." # 1 point
```

<!-- grade_id: cell-d60ffbe26d0924e8, points: 4, grade: true, solution: false -->
```python
assert len(sklearn_tokens) == 28, "Scikit-learn token list is empty." # 4 points
```

3. **Question (Markdown response)** [3 points]:  
* Compare NLTK vs. Scikit-learn tokenization

<!-- grade_id: cell-f82467dc0b20273b, points: 3, grade: true, solution: true -->
### Answer 3 

YOUR ANSWER HERE

4. Convert text to **lowercase**, **remove punctuation** (use regular expression with `re`), and remove **stopwords** using `NLTK`

<!-- grade_id: cell-6527db2cbde486ea, points: 0, grade: false, solution: true -->
```python
import re
from nltk.corpus import stopwords

nltk.download("stopwords")
stop_words = set(stopwords.words("english"))


# YOUR CODE HERE
raise NotImplementedError()

print("NLTK Cleaned:", nltk_cleaned)
```

<!-- grade_id: cell-feb0c3923a0ab884, points: 1, grade: true, solution: false -->
```python
assert isinstance(nltk_cleaned, list), "NLTK tokenization failed." # 1 point
```

<!-- grade_id: cell-1bb270dc5bd12cd2, points: 4, grade: true, solution: false -->
```python
assert len(nltk_cleaned) == 19, "NLTK token list is empty." # 4 point
```

5. **Question (Markdown response)** (4 points):  
- Contrast the NLTK cleaned result with the sklearn one from question 2? 

<!-- grade_id: cell-784db958fe07bf68, points: 4, grade: true, solution: true -->
### Answer 5

YOUR ANSWER HERE

```python

```